---
  title: "search 1"
  description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: d99522c061   
```


search in context


`@instructions`
search in instructions

`@hint`
search in hint

`@pre_exercise_code`
```{python}
# pre ex code
```
`@sample_code`
```{sql}
SELECT * FROM table
```
`@solution`
```{sql}
SELECT * FROM table
```
`@sct`
```{python}
# search sct
```
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: d5186b9848   
```

`@projector_key`
d195e836de02ed3e6ea5f515552ce6fb
---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: c7a26e307a   
```


tab exercise search


`@pre_exercise_code`
```{python}
# pre ex code - tab exercise
```
`@sample_code`
```{sql}
# sample code - tab exercise
```
***



```yaml
type: NormalExercise 
xp: 100 
key: 316ce98616   
```





`@instructions`
tab exercise - sub exercise 1

`@hint`
search hint

`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


