---
title: '5'
description: ""
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: f3b3438d42   
```

`@projector_key`
26cb05e2f92a2d96a3595298a80441ef
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 7c5c4c10e5   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: ca132924ac   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 2d7ed729dd   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`

