---
  title: "chapter 2"
  v2: true

---
## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: ea6b7175bc   
```

adas

`@instructions`
dsadasdasdsasadsa

`@hint`


`@pre_exercise_code`

```{python}

```

`@sample_code`

```{sql}

```

`@solution`

```{sql}

```

`@sct`

```{python}

```





