---
title: 'search 1'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: d99522c061
xp: 100
```

search in context

`@instructions`
search in instructions

`@hint`
search in hint

`@pre_exercise_code`
```{python}
# pre ex code
```

`@sample_code`
```{sql}
SELECT * FROM table
```

`@solution`
```{sql}
SELECT * FROM table
```

`@sct`
```{python}
# search sct
```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise
key: d5186b9848
xp: 50
```

`@projector_key`
d195e836de02ed3e6ea5f515552ce6fb

---

## Insert exercise title here

```yaml
type: TabExercise
key: c7a26e307a
xp: 100
```

tab exercise search

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}
# pre ex code - tab exercise
```

`@sample_code`
```{sql}
# sample code - tab exercise
```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


***

```yaml
type: NormalExercise
key: 316ce98616
xp: 100
```



`@instructions`
tab exercise - sub exercise 1

`@hint`
search hint

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`
