---
  title: "Template Chapter 1"
  description: "This is a template chapter."
---

## Ex 1.1

```yaml
type: NormalExercise 
lang: sql
xp: 100 
skills: 1
key: 523146ceda   
```


**ok this is my text and it's bold**

> this is a quote

![](image-url)

- bulleted list
- 2 
- # ## __****3****__


`@instructions`
dadaremoved


aaa

`@hint`


`@pre_exercise_code`

```{python}
pre exercise code
```

`@sample_code`

```{sql}
dada
```

`@solution`

```{sql}
dada
```

`@sct`

```{python}
da
```

---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 5c8ca179fc   
```

`@projector_key`
52fd7326b63833b8c981091a16c9b232
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 7499d99245   
```





---

## bla

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: fc01042ec4   
```


assignment for pure mcq!


`@hint`
this is a just simple hint

`@possible_answers`
- first
- second
- [third]
- fourth
- fiftha
- sixt

`@feedbacks`
- not
- not again
- yes!
- not again
- not again
- not again

---

## Inserted at index 5

```yaml
type: TabExercise 
xp: 100 
key: b8eed3ea64   
```





---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: e128a81496   
```


dasdsa


`@instructions`
dsadas

---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: ccef3b2bc2   
```


normal


---

## new

```yaml
type: NormalExercise 
xp: 100 
key: 403425b3af   
```


ok here


`@instructions`
instructions

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 5bdf09a717   
```





---

## name

```yaml
type: VideoExercise 
xp: 50 
key: 88e2a0b79e   
```

`@projector_key`
c5ab42ef1d11d6712d40a5dbbc611b35
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 25f3484ea7   
```


assignment1


---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: 42beafd3ed   
```


dasdas


`@sample_code`

```{sql}
dasdas
```

***



```yaml
type: NormalExercise 
xp: 20 
key: 2f6d7e120a   
```





`@instructions`


`@hint`


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: 45f055efd3   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



***



```yaml
type: NormalExercise 
xp: 20 
key: 3328871cf6   
```





`@instructions`


`@hint`


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: eaffb1db19   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: 9b3769de35   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



***



```yaml
type: NormalExercise 
xp: 0 
key: 4063e84269   
```





`@instructions`


`@hint`


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```



---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 94daf41631   
```

`@projector_key`
4f4fe85a97cb4553ccb8efed206b8a18
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 12a7d3e853   
```





---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 3216ee61c0   
```




