---
  title: "3"
  description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: dede115e92   
```





---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: c7b496ad5b   
```





---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: 2345528580   
```





***



```yaml
type: NormalExercise 
xp: 20 
key: fcb7964427   
```





`@instructions`


`@hint`


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: 14c1c56dc4   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: 19e21bafb5   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: 19e21bafb5   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: f6ae6b441a   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: d5e1d0f9a5   
```





---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: ed1334b0da   
```





***



```yaml
type: NormalExercise 
xp: 50 
key: 8fbc4239d1   
```





`@instructions`


`@hint`


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 99024fb6b2   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



---

## test

```yaml
type: TabExercise 
xp: 100 
key: acc739051f   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


***



```yaml
type: NormalExercise 
xp: 25 
key: 4b1ba7a939   
```





`@instructions`


`@hint`


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 25 
key: 2acd9ab313   
```





`@instructions`


`@hint`


`@sct`

```{python}

```



***



```yaml
type: NormalExercise 
xp: 25 
key: eca4a4997c   
```





`@instructions`


`@hint`


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```



***



```yaml
type: MultipleChoiceExercise 
xp: 25 
key: abcf8c38f3   
```





`@instructions`


`@hint`


`@sct`

```{python}

```


