---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1Ex 1.1

```yaml
type: NormalExercise 
lang: sql
xp: 100 
skills: 1
key: 523146ceda   
```


In this exercise, you will look at an equity index and plot it for a particular range of dates.
The data used in this exercise and in the rest of the course are contained in the package [qrmdata](https://www.rdocumentation.org/packages/qrmdata/versions/2016-01-03-1). You also need the package [xts](https://www.rdocumentation.org/packages/xts/versions/0.9-7) to manipulate time series.

When the `qrmdata` library is attached, as it will be throughout the course, you can load a dataset with the `data()` command. For example, the command `data("FTSE")` loads the UK FTSE (Financial Times Stock Exchange) index, which you can then refer to as object `FTSE`.

will you

If you want to extract the data from a certain date range, for example from April 1 to June 30, 2000, you can create a new object using the command `ftse00 <- FTSE["2000-04-01/2000-06-30"]`.

From now onwards, the `xts` package will also already be loaded into your workspace.


`@instructions`
dadaremoveddasdas
dsada
aaadasdsada

`@hint`


`@pre_exercise_code`

```{python}
pre exercise code
import a from b
```


`@sample_code`

```{sql}
dada
```


`@solution`

```{sql}
dada
```


`@sct`

```{python}
da
```


`@possible_answers`


`@feedback`


---

## Insert exercise title here goran goran goran

```yaml
type: VideoExercise 
xp: 50 
key: 5c8ca179fc   
```

`@projector_key`
52fd7326b63833b8c981091a16c9b232
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 7499d99245   
```


will you find this


`@instructions`


`@hint`
dasdas

`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## bla

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: fc01042ec4   
```


will you find this


`@instructions`


`@hint`
aathis is a just simple hint

`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`
- first
- second
- [third]
- fourth
- fiftha
- sixt

`@feedback`
- not
- not again
- yes!
- not again
- not again
- not again

---

## Inserted at index 5

```yaml
type: TabExercise 
xp: 100 
key: b8eed3ea64   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: e128a81496   
```


will you find this


`@instructions`
dsadas

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: ccef3b2bc2   
```


normal


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## new

```yaml
type: NormalExercise 
xp: 100 
key: 403425b3af   
```


ok here


`@instructions`
instructionsa

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 5bdf09a717   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## name

```yaml
type: VideoExercise 
xp: 50 
key: 88e2a0b79e   
```

`@projector_key`
c5ab42ef1d11d6712d40a5dbbc611b35
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 25f3484ea7   
```


assignment1


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: 42beafd3ed   
```


dasdsa


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}
dasdas
```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


***



```yaml
type: NormalExercise 
xp: 20 
key: 2f6d7e120a   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: 45f055efd3   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`



***



```yaml
type: NormalExercise 
xp: 20 
key: 3328871cf6   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: eaffb1db19   
```





`@instructions`
subexercise

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`



***



```yaml
type: MultipleChoiceExercise 
xp: 20 
key: 9b3769de35   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`



***



```yaml
type: NormalExercise 
xp: 0 
key: 4063e84269   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`



---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 94daf41631   
```

`@projector_key`
4f4fe85a97cb4553ccb8efed206b8a18
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 12a7d3e853   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 3216ee61c0   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 0ebffc1c99   
```

`@projector_key`
3da3f2447ac05eb53124f4639f60eaa2
---

## mp3

```yaml
type: VideoExercise 
xp: 50 
key: 9e5dad7c00   
```

`@projector_key`
d01e39a5b316ef63e7fe717882fbf199
