---
title: Template Chapter 1
description: >-
  This is a template chapter.


---
## Ex 1.1

```yaml
type: NormalExercise
lang: sql
xp: 100
skills: 1
key: 523146ceda
```

# ok this is a hader
das
das
d

dasdas


**bolded**

`@instructions`
removed


aaa


`@pre_exercise_code`
```{python}
pre exercise code
```
`@sample_code`
```{sql}
dada
```
`@solution`
```{sql}
dada
```
`@sct`
```{python}
da
```





---
## Insert exercise title here

```yaml
type: VideoExercise

xp: NaN

key: 5c8ca179fc
```

`@projector_key`
52fd7326b63833b8c981091a16c9b232

---
## Insert exercise title here

```yaml
type: NormalExercise

xp: NaN

key: 7499d99245
```














---
## pure mcq

```yaml
type: PureMultipleChoiceExercise

xp: NaN

key: fc01042ec4
```

assignment for pure mcq!


`@hint`
this is a just simple hint





`@possible_answers`
- first
- second
- [third]
- fourth
- fiftha
- sixt

`@feedbacks`
- not
- not again
- yes!
- not again
- not again
- not again




---
## Insert exercise title here

```yaml
type: NormalExercise

xp: NaN

key: e128a81496
```

dasdsa

`@instructions`
dsadas











---
## Insert exercise title here

```yaml
type: NormalExercise

xp: NaN

key: ccef3b2bc2
```

normal












---
## new

```yaml
type: NormalExercise

xp: NaN

key: 403425b3af
```

ok here

`@instructions`
instructions











---
## Insert exercise title here

```yaml
type: VideoExercise

xp: undefined

key: 545527377e
```

`@projector_key`
729fda1002de6560737138bed1509ebb

---
## Insert exercise title here

```yaml
type: NormalExercise

xp: undefined

key: d374ab6e91
```













