---
title: 'chapter 2'
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: ea6b7175bc   
```


changes done


`@instructions`
dsadasdasdsasadsa

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


