---
title: '3'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: d5e1d0f9a5
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: dede115e92
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: c7b496ad5b
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: TabExercise
key: 2345528580
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


***

```yaml
type: NormalExercise
key: fcb7964427
xp: 20
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: 14c1c56dc4
xp: 20
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: 19e21bafb5
xp: 20
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: 19e21bafb5
xp: 20
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: f6ae6b441a
xp: 20
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

---

## Insert exercise title here

```yaml
type: TabExercise
key: ed1334b0da
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


***

```yaml
type: NormalExercise
key: 8fbc4239d1
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: 99024fb6b2
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

---

## test

```yaml
type: TabExercise
key: acc739051f
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


***

```yaml
type: NormalExercise
key: 4b1ba7a939
xp: 20
```



`@instructions`
1

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: 2acd9ab313
xp: 20
```



`@instructions`
2

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: NormalExercise
key: eca4a4997c
xp: 20
```



`@instructions`
3

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: abcf8c38f3
xp: 20
```



`@instructions`
4

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: MultipleChoiceExercise
key: 63be9840d6
xp: 20
```



`@instructions`
exercise 5

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

***

```yaml
type: NormalExercise
key: d5a78be106
```



`@instructions`
6

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`

---

## Insert exercise title here

```yaml
type: VideoExercise
key: ca7cefc136
xp: 50
```

`@projector_key`
9fbe455dcb1a900783c18299dd942496

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 2a6ed7882f
xp: 50
```

`@projector_key`
d2c2c2f686b5a33263aaa72957527204
