---
title: '5'
description: ""
---

## Insert exercise title here

```yaml
type: VideoExercise
key: f3b3438d42
xp: 50
```

`@projector_key`
26cb05e2f92a2d96a3595298a80441ef

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 7c5c4c10e5
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: ca132924ac
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 2d7ed729dd
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`
